---
title: "MA22004 Lab 4"
subtitle: "Confidence Intervals"
description: "Lab report consists of answers to exercises 1-5."
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
tutorial:
  id: "ma22004-lab04"
  version: 1.0.1020
---

```{css, echo=FALSE}
body {
    counter-reset: li; /* initialize counter named li */
}

ol {
    margin-left:0; /* Remove the default left margin */
    padding-left:0; /* Remove the default left padding */
}
ol > li {
    position:relative; /* Create a positioning context */
    margin:0 0 10px 2em; /* Give each list item a left margin to make room for the numbers */
    padding:10px 80px; /* Add some spacing around the content */
    list-style:none; /* Disable the normal item numbering */
    border-top:2px solid #317EAC;
    background:rgba(49, 126, 172, 0.1);
}
ol > li:before {
    content:"Exercise " counter(li); /* Use the counter as content */
    counter-increment:li; /* Increment the counter by 1 */
    /* Position and style the number */
    position:absolute;
    top:-2px;
    left:-2em;
    -moz-box-sizing:border-box;
    -webkit-box-sizing:border-box;
    box-sizing:border-box;
    width:7em;
    /* Some space between the number and the content in browsers that support
       generated content but not positioning it (Camino 2 is one example) */
    margin-right:8px;
    padding:4px;
    border-top:2px solid #317EAC;
    color:#fff;
    background:#317EAC;
    font-weight:bold;
    font-family:"Helvetica Neue", Arial, sans-serif;
    text-align:center;
}
li ol,
li ul {margin-top:6px;}
ol ol li:last-child {margin-bottom:0;}

.oyo ul {
    list-style-type:decimal;
}

.oyo ul {
    list-style-type:decimal;
}


hr {
    border: 1px solid #357FAA;
}

div#boxedtext {
    background-color: rgba(86, 155, 189, 0.2);
    padding: 20px;
    margin-bottom: 20px;
    font-size: 10pt;
}

div#template {
    margin-top: 30px;
    margin-bottom: 30px;
    color: #808080;
    border:1px solid #808080;
    padding: 10px 10px;
    background-color: rgba(128, 128, 128, 0.2);
    border-radius: 5px;
}

div#license { 
    margin-top: 30px;
    margin-bottom: 30px;
    color: #4C721D;
    border:1px solid #4C721D;
    padding: 10px 10px;
    background-color: rgba(76, 114, 29, 0.2);
    border-radius: 5px;
}

/*------------ Fancy blocks --------------*/

.noteblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #042e6b;
  background: 1.5em center/2.5em no-repeat;
  background-image: url("images/info-circle.svg");
}

.tipblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #d5d5ce;
  background: 1.5em center/2em no-repeat;
  background-image: url("images/lightbulb.svg");
}

.warningblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #be4e00;
  background: 1.5em center/2.5em no-repeat;
  background-image: url("images/exclamation-triangle.svg");
}

.importantblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #c30000;
  background: 1.5em center/2.5em no-repeat;
  background-image: url("images/exclamation-circle.svg");
}

.exerciseblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #317EAC;
  background: rgba(49, 126, 172, 0.1) 1.5em center/2.5em no-repeat;
  background-image: url("images/flag.svg");
}
```

```{r setup, include = FALSE}
# packages
library(MA22004labs)
library(tidyverse)
library(infer)
library(learnr)
library(gradethis)
library(knitr)

# tutorial options
tutorial_options(
  # code running in exercise times out after 30 seconds
  # if it is taking more than 30 s something is wrong 
  exercise.timelimit = 30,
  # use gradethis for checking
  exercise.checker = gradethis::grade_learnr
  )
options(width = 60)

# hide non-exercise code chunks
knitr::opts_chunk$set(echo = FALSE, error = TRUE, 
                      comment = "", fig.align = "center", fig.width = 6, 
                      fig.asp = 0.75, out.width = "100%")

data(ames)
population <- ames
samp <- sample_n(ames, 60)
```

## Getting started

If you have access to data on an entire population, say the size of every house in Ames, Iowa, it's straight forward to answer questions like, "How big is the typical house in Ames?" and "How much variation is there in sizes of houses?". If you have access to only a sample of the population, as is often the case, the task becomes more complicated. What is your best guess for the typical size if you only know the sizes of several dozen houses? This sort of situation requires that you use your sample to make inference on what your population looks like.

:::{.noteblock}
The report for Lab 4 consists of answers to exercises 1-5 in the last section of this tutorial. Upload the submission to [Gradescope](https://www.gradescope.com) as a PDF; an `rmarkdown` template called *MA22004 Lab Report* is provided. Further instructions on using the class template to produce a nice lab report containing both data analysis and text are in Section "RStudio and lab reports" of Lab 1. 
:::

### Packages 

In this lab, we will explore and visualize the data using the **tidyverse** suite of packages, and perform statistical inference using **infer**.

Recall that all the necessary packages can be loaded using the `library` command. 

```{r load-packages, echo = TRUE, eval = FALSE}
library(MA22004labs)
library(tidyverse)
library(infer)
```

### The data

:::{.tipblock}
In the previous lab, "Lab 3: Sampling Distributions", we looked at the **Ames Housing Data**, the population data
of houses from Ames, Iowa. Recall that this data can be loaded by calling:

```{r load-data, echo = TRUE, eval = FALSE}
data(ames)
```
:::

In this lab we'll start with a simple random sample of size 60 from the population. Specifically, this is a random sample of 60 row (i.e. all the response variables) from the data frame `ames` (which we save as `population` for pedagogical purposes). Note that the data set has information on many housing variables, but for the first portion of the lab we'll focus on the size of the house, represented by the variable `area`.

```{r sample, echo = TRUE, eval = FALSE}
population <- ames
samp <- sample_n(population, 60)
```

Explore the distribution for the population and the sample by plotting histograms and creating summary tables. 

```{r sample-exploration, exercise = TRUE}

```

```{r sample-exploration-hint}
summary(`___`)
ggplot(`___`) + geom_histogram()
```

```{r sample-exploration-solution}
summary(samp$area)
ggplot(samp, aes(x = area)) + geom_histogram()
```

:::{.exerciseblock}
Describe the distribution of your sample `samp$area`. What would you say is the "typical" size within your sample? Also state precisely what you interpreted "typical" to mean (e.g., reference a summary statistic).

```{r answer-1a, exercise = TRUE}

```

```{r answer-1b}
question_text(
  "Enter your thoughts and observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::

:::{.exerciseblock}
Would you expect another student's distribution for `samp$area` to be identical to yours? Would you expect it to be similar? Why or why not?

```{r answer-2}
question_text(
  "Enter your thoughts and observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::

## Confidence intervals

One of the most common ways to describe the typical or central value of a distribution is to use the mean. Calculate the mean of the sample area using the `mean` function.

```{r sample-mean, exercise = TRUE, exercise.warn_invisible = FALSE}

```

```{r sample-mean-hint}
mean(`___`)
```

```{r sample-mean-solution}
mean(samp$area)
```

```{r sample-mean-check}
grade_result(
  pass_if(~ .result[1] == mean(samp$area))
)
```

```{r sample-mean-e, echo = TRUE, eval = TRUE}
sample_mean <- mean(samp$area)
```

Return for a moment to the question that first motivated this lab: based on this sample, what can we infer about the population? Based only on this single sample, the best estimate of the average living area of houses sold in Ames would be the sample mean, usually denoted as $\bar{x}$ (here we're calling it `sample_mean`). That serves as a good *point estimate* but it would be useful to also communicate how uncertain we are of that estimate. This can be captured by using a *confidence interval*.

### Normality theory

One way of calculating a confidence interval for a parameter is based on the Central Limit Theorem (normality theory). More precisely, we can calculate a 95% confidence interval for a population mean by adding and subtracting 1.96 standard errors to the point estimate for the sample mean (see [Section 3.1.2 of the Course Notes](https://dundeemath.github.io/MA22004/inference-single-sample.html#mean-large-sample))

\[\overline{x} \pm 1.96 \cdot \mathsf{se}\,. \] 

Compute the standard error `se` for the 60 samples and save it to a vector `se`. Also, compute upper and lower limits, `lower` and `upper`, respectively, based on 1.96 deviations of the standard error from the sample_mean. Print the confidence interval using the vector command `c(lower, upper)`. 

```{r ci, exercise = TRUE, exercise.setup = "sample-mean-e"}

```

```{r ci-hint-1}
se <- `___`
lower <- `___`
upper <- `___`
c(lower, upper)
```

```{r ci-hint-2}
se <- sd(`___`) / sqrt(`___`)
lower <- `___` - 1.96 * `___`
upper <- `___` + 1.96 * `___`
c(lower, upper)
```

```{r ci-check}
grade_result(
  pass_if(~ all( signif(.result[1], digits = 8) == signif(mean(samp$area) - 1.96 * sd(samp$area) / sqrt(60), digits = 8) & signif(.result[2], digits = 8) == signif(mean(samp$area) + 1.96 * sd(samp$area) / sqrt(60), digits = 8))
))
```

```{r ci-solution}
se <- sd(samp$area) / sqrt(60)
lower <- sample_mean - 1.96 * se
upper <- sample_mean + 1.96 * se
c(lower, upper)
```

This is an important inference that we've just made: even though we don't know what the full population looks like, we're 95% confident that the true average size of houses in Ames lies between the values *lower* and *upper*. There are a few conditions that must be met for this interval to be valid.

:::{.exerciseblock}
For the confidence interval \[\overline{x} \pm 1.96 \frac{s}{\sqrt{m}}\] computed in Section "Confidence intervals" above to be valid, the sample mean must be normally distributed and have standard error $s / \sqrt{m}$ where $s$ is the sample standard deviation and $m$ is the sample size. What conditions must be met for this to be true?

```{r answer-3}
question_text(
  "Enter your thoughts and observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::

### Bootstrapping

Another way of calculating a confidence interval for a parameter is using simulation, or to be more specific, using **bootstrapping**. The term bootstrapping comes from the phrase "pulling oneself up by one's bootstraps", which is a metaphor for accomplishing an impossible task without any outside help. In this case the impossible task is estimating a population parameter (e.g., the unknown population mean), and we'll accomplish it using data from only the given sample. Note that this notion of saying something about a population parameter using only information from an observed sample is the crux of statistical inference, it is not limited to bootstrapping.

In essence, bootstrapping assumes that there are more observations in the populations like the ones in the observed sample. So we "reconstruct" the population by resampling from our sample, with replacement. The bootstrapping scheme is as follows:

-   *Step 1.* Take a bootstrap sample - a random sample taken **with replacement** from the original sample, of the same size as the original sample.
-   *Step 2.* Calculate the bootstrap statistic - a statistic such as mean, median, proportion, slope, etc. computed on the bootstrap samples.
-   *Step 3.* Repeat steps (1) and (2) many times to create a bootstrap distribution - a distribution of bootstrap statistics.
-   *Step 4.* Calculate the bounds of the XX% confidence interval as the central XX% of the bootstrap distribution.

Instead of coding up each of these steps, one can use the **infer** package to construct confidence intervals.

Below is an overview of the functions we will use to construct this confidence interval:

+-------------+--------------------------------------------------------------------------------------------------------------------------------------------+
| Function    | Purpose                                                                                                                                    |
+:============+:===========================================================================================================================================+
| `specify`   | Identify your variable of interest                                                                                                         |
+-------------+--------------------------------------------------------------------------------------------------------------------------------------------+
| `generate`  | The number of samples you want to generate                                                                                                 |
+-------------+--------------------------------------------------------------------------------------------------------------------------------------------+
| `calculate` | The sample statistic you want to do inference with, or you can also think of this as the population parameter you want to do inference for |
+-------------+--------------------------------------------------------------------------------------------------------------------------------------------+
| `get_ci`    | Find the confidence interval                                                                                                               |
+-------------+--------------------------------------------------------------------------------------------------------------------------------------------+

This code will find the 95% confidence interval for .

```{r confidence-interval-infer, exercise = TRUE}
samp %>%
 specify(response = area) %>%
 generate(reps = 1000, type = "bootstrap") %>%
 calculate(stat = "mean") %>%
 get_ci(level = 0.95)
```

-   In `specify` we specify the `response` variable (if the statistic of interest is a proportion, we would also give the level of that variable we are calling a `success` here).
-   In `generate` we provide the number of resamples we want from the population in the `reps` argument (this should be a reasonably large number) as well as the type of resampling we want to do, which is `"bootstrap"` in the case of constructing a confidence interval.
-   Then, we `calculate` the sample statistic of interest for each of these resamples, which is "mean" in this case (but could also be "median", "prop", "sd", "diff in means", ...).

Feel free to test out the rest of the arguments for these functions, since these commands will be used together to calculate confidence intervals and solve inference problems for the rest of the semester. But we will also walk you through more examples in future chapters.

:::{.exerciseblock}
Compare the confidence interval for the true mean area produced by bootstrapping (using the `infer` package) and the normality theory. 

```{r thoughts-bootstrapping-v-normality}
question_text(
  "Enter your thoughts as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::

:::{.tipblock}
To recap: we have produced confident intervals for the true mean area of the **Ames Housing Data** using two different methods: (1) normality theory and (2) bootstrapping. If we didn't already know that the true population mean, then we would be 95% confident that the true population parameter lies within the computed interval estimates. 
:::


## Confidence levels

In this case we have the luxury of knowing the true population mean since we have data on the entire population. This value can be calculated using the following command:

```{r pop-mean, eval = TRUE, echo = TRUE}
mean(population$area)
```

:::{.exerciseblock}
What does "95% confidence" mean in relation to a confidence interval? If you're not sure, see [Section 2.2 of the Course Notes](https://dundeemath.github.io/MA22004/statistical-inference.html#confidence-intervals). Does your confidence interval capture the true average size of houses in Ames?

```{r answer-4}
question_text(
  "Enter your thoughts and observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::

Using `R`, we're going to recreate many samples to learn more about how sample means and confidence intervals vary from one sample to another. *Loops* come in handy here.

Here is the rough outline:

-   Obtain a random sample.
-   Calculate and store the sample's mean and standard deviation.
-   Repeat the first two steps 50 times.
-   Use these stored statistics to calculate many confidence intervals.

Create a loop that calculates the means and standard deviations of 50 random samples. Save the means to a vector called `samp_mean` and the standard deviations to `samp_sd`. After the loop, Construct a vector of lower limits called `lower_vector` and upper limits called `upper_vector` based on the large-sample 95% confidence intervals for the population mean. 

```{r cififty, exercise = TRUE}
samp_mean <- rep(NA, 50)   # create empty vector for means
samp_sd <- rep(NA, 50)     # create empty vector for sds
m <- 60                    # sample size

for(i in 1:50){
`___`
}

`___`
`___`

summary(lower_vector)
summary(upper_vector)
```

```{r cififty-hint-1}
for(i in 1:50){
  x <- `___`                      # obtain a sample of size 60 from the population
  samp_mean[i] <- `___`           # save sample mean in ith element of samp_mean
  samp_sd[i] <- `___`             # save sample sd in ith element of samp_sd
}

lower_vector <- `___`  
upper_vector <- `___` 
```

```{r cififty-hint-2}
for(i in 1:50){
  x <- sample(`___`, m)          # obtain a sample of size m from the population
  samp_mean[i] <- mean(`___`)    # save sample mean in ith element of samp_mean
  samp_sd[i] <- sd(`___`)        # save sample sd in ith element of samp_sd
}

lower_vector <- `___` - 1.96 * `___` / `___`
upper_vector <- `___` + 1.96 * `___` / `___`

summary(lower_vector)
summary(upper_vector)
```

```{r cififty-solution}
samp_mean <- rep(NA, 50)   # create empty vector for means
samp_sd <- rep(NA, 50)     # create empty vector for sds
m <- 60                    # sample size

for(i in 1:50){
  x <- sample(population$area, m) # obtain a sample of size m from the population
  samp_mean[i] <- mean(x)    # save sample mean in ith element of samp_mean
  samp_sd[i] <- sd(x)        # save sample sd in ith element of samp_sd
}

lower_vector <- samp_mean - 1.96 * samp_sd / sqrt(m) 
upper_vector <- samp_mean + 1.96 * samp_sd / sqrt(m) 

summary(lower_vector)
summary(upper_vector)
```


```{r cififty-e}
samp_mean <- rep(NA, 50)   # create empty vector for means
samp_sd <- rep(NA, 50)     # create empty vector for sds
m <- 60                    # sample size
for(i in 1:50){
  x <- sample(population$area, m) # obtain a sample of size m = 60 from the population
  samp_mean[i] <- mean(x)    # save sample mean in ith element of samp_mean
  samp_sd[i] <- sd(x)        # save sample sd in ith element of samp_sd
}

lower_vector <- samp_mean - 1.96 * samp_sd / sqrt(m) 
upper_vector <- samp_mean + 1.96 * samp_sd / sqrt(m) 
```

Included with the lab is a non-standard function called `plot_ci`; use the command `plot_ci(lower_vector, upper_vector, mean(population))` to plot fifty (50) 95% confidence intervals for the true mean area from the **Ames House Data**. This tutorial is set-up to rerun the loop above in the background every time you call the code chunk below; try rerunning the code.

```{r plot-fifty-cis, exercise = TRUE, exercise.setup = "cififty-e"}
plot_ci(lower_vector, upper_vector, mean(population$area))
```

:::{.exerciseblock}
Included with the lab is a non-standard function called `plot_ci`; use the command `plot_ci(lower_vector, upper_vector, mean(population))` to plot fifty (50) 95% confidence intervals for the true mean area from the **Ames House Data**. What proportion of your confidence intervals include the population mean? Is this proportion exactly equal to the confidence level? If not, explain why.

```{r answer-5a, exercise = TRUE, exercise.setup = "cififty-e"}

```

```{r answer-5b}
question_text(
  "Enter your thoughts and observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::


## Summary (Exercises 1-5)

:::{.warningblock}
To submit your completed laboratory report, submit the answers (only) to the exercises 1-5 below. Upload the submission to [Gradescope](https://www.gradescope.com) as a PDF. To create your new lab report, in RStudio, go to **New File \> R Markdown** then choose **From Template** and select *MA22004 Lab Report* from the list of templates. Further instructions on using the class template to produce a nice lab report containing both data analysis and text are in the first Section "RStudio and lab reports" of Lab 1.
:::

:::{.importantblock}
As this lab requires you to take random samples to build sampling distributions, fix a random seed value in your `load-packages` chunk:

```{r setseed, echo = TRUE, eval = FALSE}
set.seed(n)
```

where you replace `n` by your favorite integer. 
:::

The following exercises refer to variables found in or derived from the **Ames Housing Data**  in the data frame `ames` which is packaged with `MA22004labs`. In particular, we focus on two features: the above ground living area of the house in square feet (`area` i.e. `ames$area`) and the sale price (`price` i.e. `ames$price`).  

1.  Describe the distribution of your sample `samp <- sample(ames$area, 60)`. What would you say is the "typical" size within your sample? Also state precisely what you interpreted "typical" to mean (e.g., reference a summary statistic). **[3 marks]**

2.  Would you expect another student's distribution for `samp <- sample(ames$area, 60)` to be identical to yours? Would you expect it to be similar? Why or why not? **[3 mark]**

3.  For the confidence interval \[\overline{x} \pm 1.96 \frac{s}{\sqrt{m}}\] computed above in Section "Confidence intervals" to be valid the sample mean must be normally distributed and have standard error, $s / \sqrt{m}$, where $s$ is the sample standard deviation and $m$ is the sample size. What conditions must be met for this to be true? **[3 marks]**

<!-- 4.  What does "95% confidence" mean in relation to a confidence interval? If you're not sure, see [Section 2.2 of the Course Notes](https://dundeemath.github.io/MA22004/statistical-inference.html#confidence-intervals). **[3 marks]** -->

4.  Included with the lab is a non-standard function called `plot_ci`; use the command `plot_ci(lower_vector, upper_vector, mean(population))` to plot fifty 95% confidence intervals for the true mean area from the **Ames House Data**. What proportion of your confidence intervals include the population mean? Is this proportion exactly equal to the confidence level? If not, explain why. **[5 marks]**

<!-- 5.  Pick a confidence level of your choosing, provided it is not 95%. What is the appropriate critical value? **[2 mark]** -->

<!-- 5.  Calculate fifty new confidence intervals for the true mean area at the confidence level different from 95%. Using the `plot_ci` function, plot all intervals and calculate the proportion of intervals that include the true population mean. How does this percentage compare to the confidence level selected for the intervals? **[6 marks]** -->

5.  Using code from the **infer** package and data from the one sample of size 60 from `ames`, find a confidence interval for the proportion of Ames residents that live in single family home (`Bldg.Type == 1Fam`) with a confidence level of your choosing (other than 95%) and interpret it. **[6 marks]**


------------------------------------------------------------------------

<a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">![Creative Commons License](https://i.creativecommons.org/l/by-sa/4.0/88x31.png){style="border-width:0"}</a><br />This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>.

This work is derivative of OpenIntro Labs and was modified by Eric Hall (2020). The original labs were adapted for OpenIntro by Andrew Bray and Mine Çetinkaya-Rundel from labs written by Mark Hansen of UCLA Statistics.

