---
title: "Lab 2"
subtitle: "Intro to Data Analysis"
output:
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
tutorial:
  id: "ma22004-lab02"
  version: 0.5.2020
---

```{css, echo=FALSE}
body {
    counter-reset: li; /* initialize counter named li */
}

ol {
    margin-left:0; /* Remove the default left margin */
    padding-left:0; /* Remove the default left padding */
}
ol > li {
    position:relative; /* Create a positioning context */
    margin:0 0 10px 2em; /* Give each list item a left margin to make room for the numbers */
    padding:10px 80px; /* Add some spacing around the content */
    list-style:none; /* Disable the normal item numbering */
    border-top:2px solid #317EAC;
    background:rgba(49, 126, 172, 0.1);
}
ol > li:before {
    content:"Exercise " counter(li); /* Use the counter as content */
    counter-increment:li; /* Increment the counter by 1 */
    /* Position and style the number */
    position:absolute;
    top:-2px;
    left:-2em;
    -moz-box-sizing:border-box;
    -webkit-box-sizing:border-box;
    box-sizing:border-box;
    width:7em;
    /* Some space between the number and the content in browsers that support
       generated content but not positioning it (Camino 2 is one example) */
    margin-right:8px;
    padding:4px;
    border-top:2px solid #317EAC;
    color:#fff;
    background:#317EAC;
    font-weight:bold;
    font-family:"Helvetica Neue", Arial, sans-serif;
    text-align:center;
}
li ol,
li ul {margin-top:6px;}
ol ol li:last-child {margin-bottom:0;}

.oyo ul {
    list-style-type:decimal;
}

.oyo ul {
    list-style-type:decimal;
}


hr {
    border: 1px solid #357FAA;
}

div#boxedtext {
    background-color: rgba(86, 155, 189, 0.2);
    padding: 20px;
    margin-bottom: 20px;
    font-size: 10pt;
}

div#template {
    margin-top: 30px;
    margin-bottom: 30px;
    color: #808080;
    border:1px solid #808080;
    padding: 10px 10px;
    background-color: rgba(128, 128, 128, 0.2);
    border-radius: 5px;
}

div#license { 
    margin-top: 30px;
    margin-bottom: 30px;
    color: #4C721D;
    border:1px solid #4C721D;
    padding: 10px 10px;
    background-color: rgba(76, 114, 29, 0.2);
    border-radius: 5px;
}

/*------------ Fancy blocks --------------*/

.noteblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #042e6b;
  background: 1.5em center/2.5em no-repeat;
  background-image: url("images/info-circle.svg");
}

.tipblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #d5d5ce;
  background: 1.5em center/2em no-repeat;
  background-image: url("images/lightbulb.svg");
}

.warningblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #be4e00;
  background: 1.5em center/2.5em no-repeat;
  background-image: url("images/exclamation-triangle.svg");
}

.importantblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #c30000;
  background: 1.5em center/2.5em no-repeat;
  background-image: url("images/exclamation-circle.svg");
}

.exerciseblock {
  padding: 1em 1em 1em 6em;
  margin-bottom: 20px;
  margin-top: 20px;
  border-left: 6px solid #317EAC;
  background: rgba(49, 126, 172, 0.1) 1.5em center/2.5em no-repeat;
  background-image: url("images/flag.svg");
}
```

```{r setup, include = FALSE}
# packages
library(MA22004labs)
library(tidyverse)
library(ggmosaic)
library(janitor)
library(learnr)
library(gradethis)
library(knitr)

# tutorial options
tutorial_options(
  # code running in exercise times out after 30 seconds
  # if it is taking more than 30 s something is wrong 
  exercise.timelimit = 30,
  # use gradethis for checking
  exercise.checker = gradethis::grade_learnr
  )
options(width = 60)

# hide non-exercise code chunks
knitr::opts_chunk$set(echo = FALSE, error = TRUE, 
                      comment = "", fig.align = "center", fig.width = 6, 
                      fig.asp = 0.75, out.width = "100%")
data(cdc)
```


## Getting Started

Some define Statistics as the field that focuses on turning information into knowledge.  The first step in that process is to summarize and describe the raw information - the data. In this lab, you will gain insight into public health by generating simple graphical and numerical summaries of a data set collected by the Centers for Disease Control and Prevention (CDC). As this is a large data set, along the way you'll also learn the indispensable skills of data processing and subsetting.

:::{.noteblock}
The report for Lab 2 consists of answers to exercises 1-6 in the last section of this tutorial. Upload the submission to [Gradescope](https://www.gradescope.com) as a PDF. To create your new lab report, in RStudio, go to **New File \> R Markdown** then choose **From Template** and select *MA22004 Lab Report* from the list of templates. Further instructions on using the class template to produce a nice lab report containing both data analysis and text are in the first Section "RStudio and lab reports" of Lab 1.
:::

### Packages 

In this lab, we will explore and visualize the data using the **tidyverse** suite of packages.  The package **ggmosaic** provides the `geom_mosaic` geometry for producing mosaic plots with `ggplot`. The package **janitor** is for data cleaning and provides the function `tabyl` for summarizing data. 

Recall that all the necessary packages can be loaded using the `library` command. 

```{r load-packages, echo = TRUE, eval = FALSE}
library(MA22004labs)
library(tidyverse)
library(ggmosaic)
library(janitor)
```


### The **CDC Health Data**

:::{.tipblock}
This lab will look at the **CDC Health Data** which has been loaded in this tutorial under the data frame `cdc`. 

The Behavioral Risk Factor Surveillance System (BRFSS) is an annual telephone survey of 350,000 people in the United States collected by the CDC. As its name implies, the BRFSS is designed to identify risk factors in the adult population and report emerging health trends. For example, respondents are asked about their diet and weekly physical activity, their HIV/AIDS status, possible tobacco use, and even their level of healthcare coverage. The BRFSS Web site ([http://www.cdc.gov/brfss](http://www.cdc.gov/brfss)) contains a complete description of the survey, including the research questions that motivate the study and many interesting results derived from the data.

This data set contains a random sample of 20,000 people from the BRFSS survey conducted in 2000. While there are over 200  variables in this data set, we will work with a small subset. The data set of 20,000 observations can be loaded using the `data` command. 

```{r load-cdc-data, eval = FALSE, echo = TRUE}
data(cdc)
```
:::

The data set `cdc` that shows up in your workspace is a *data matrix*, with each row representing a *case* and each column representing a *variable*.  `R` calls this data format a *data frame*, which is a term that will be used throughout the labs.

To view the names of the variables, one uses the `names` function.

```{r names, exercise = TRUE}
names(``___``)
```

This returns the names `genhlth`, `exerany`, `hlthplan`, `smoke100`, `height`, `weight`, `wtdesire`, `age`, and `gender`. Each one of these variables corresponds to a question that was asked in the survey. For example, for `genhlth`, respondents were asked to evaluate their general health, responding either excellent, very good, good, fair or poor. The `exerany` variable indicates whether the respondent exercised in the past month (1) or did not (0). Likewise, `hlthplan` indicates whether the respondent had some form of health coverage (1) or did not (0). The `smoke100` variable indicates whether the respondent had smoked at least 100 cigarettes in her lifetime. The other variables record the respondent's `height` in inches, `weight` in pounds as well as their desired weight, `wtdesire`, `age` in years, and `gender`.

We can have a look at the first few entries (rows) of our data with the `head` function.

```{r head, exercise = TRUE}
`___`(cdc)
```

Similarly we can look at the last few using the `tail` function.

```{r tail, exercise = TRUE}
`___`(cdc)
```

Remember that you can use `glimpse` to take a quick peek at your data to understand its contents better.

You could also look at *all* of the data frame at once by typing its name into the console, but that might be unwise here.  We know `cdc` has 20,000 rows, so viewing the entire data set would mean flooding your screen.  It's better to take small peeks at the data with `head`, `tail` or the subsetting techniques that you'll learn in a moment.

## Summary stats and tabyls

The BRFSS questionnaire is a massive trove of information.  A good first step in any analysis is to distill all of that information into a few summary statistics and graphics.  As a simple example, the function `summary` returns a numerical summary: minimum, first quartile, median, mean, second quartile, and maximum. 

Try producing a summary for the `weight` variable.

```{r summary-weight, exercise = TRUE}
summary(`___`$`___`)
```

If you wanted to compute the interquartile range for the respondents' weight, you would look at the output from the summary command above and then enter:

```{r weight-range-arith, eval = FALSE, echo = TRUE}
190 - 140
```

`R` also has built-in functions to compute summary statistics one by one. For instance, the mean, median, and variance of `weight` can be calculated using the obvious functions.

:::{.noteblock}
**Summary statistics:** Some useful function calls for summary statistics for a single numerical variable are as follows:

-   `mean`
-   `median`
-   `sd`
-   `var`
-   `IQR`
-   `min`
-   `max`

Note that each of these functions takes a single vector as an argument and returns a single value.
:::

While it makes sense to describe a quantitative variable like `weight` in terms of these statistics, what about categorical data?  We would instead consider the sample frequency or relative frequency distribution.  The function `tabyl` does this for you by counting the number of times each kind of response was given. 

Use the `tabyl` function to see the number of people who have smoked 100 cigarettes in their lifetime (`smoke100`).

```{r table-smoke, exercise = TRUE}
cdc %>% tabyl(smoke100)
```

The output of the `tabyl` function includes both the counts and the relative frequency distribution for each category of response to the question `smoke100` (in this case yes `1` or no `0`).

Next, make a bar plot of the entries in the `tabyl` data frame using `ggplot` with the `geom_bar` geometry. Here we must set the option `stat = "identity"` within the geometry (the default is `stat = count` but we've already summarized the data using `tabyl`; there are other ways to produce the plot below...).

```{r table-smoke-barplot, exercise = TRUE}
df <- cdc %>% tabyl(smoke100)
ggplot(`___`, aes(x = smoke100, y = `___`)) + geom_bar(stat = "identity")
```

```{r table-smoke-barplot-hint}
df <- cdc %>% tabyl(smoke100)
ggplot(df, aes(x = smoke100, y = `___`)) + geom_bar(stat = "identity") 
# you can set `y` to `n` or `percentage` as you please
```

Notice what we've done here! We've computed a summary table of the responses to `cdc$smoke100`.

:::{.exerciseblock}
Using the **CDC Data**, create a numerical summary for `height` and `age`, and compute the interquartile range for each. Compute the relative frequency distribution for `gender` using `tabyl`. How many males are in the sample? What proportion of the sample reports being in excellent health? 

```{r answer-4, exercise = TRUE}

```

```{r answer-4-solution}
# The following commands can aid in answering the questions
summary(cdc$height)
summary(cdc$age)
IQR(cdc$height)
IQR(cdc$age)

cdc %>% tabyl(gender)
cdc %>% tabyl(genhlth)
```

```{r answer-4-free}
question_text(
  "Enter your observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::

The `tabyl` command can be used to tabulate any number of variables that you provide. For example, to examine which participants have smoked across each gender, we could use the following.

```{r table-smoke-gender, eval = FALSE, echo = TRUE}
cdc %>% tabyl(smoke100, gender)
```

Here, we see column labels of 0 and 1. Recall that 1 indicates a respondent has smoked at least 100 cigarettes. The rows refer to gender (m for male or f for female). 

Creating a mosaic plot of this table (`gender` vs `smoke100`) can be done using `ggplot` and working directly from the original `cdc` data frame by setting the aesthetic for `x` in `geom_mosaic` to `product(X,Y)`. 

```{r mosaic-smoke-gender, exercise = TRUE}
ggplot(cdc) + geom_mosaic(aes(x = product(smoke100, gender)))
```

:::{.exerciseblock}
Generate a mosaic plot of `gender` vs `smoke100` from the **CDC Health Data**. What does the mosaic plot of reveal about smoking habits and gender? 

```{r answer-5a, exercise = TRUE}

```

```{r answer-5b}
question_text(
  "Enter your observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```
:::

## Interlude: How `R` thinks about data

We mentioned that `R` stores data in data frames, which you might think of as a type of spreadsheet. Each row is a different observation (a different respondent) and each column is a different variable (the first is `genhlth`, the second `exerany` and so on). Use the `dim` function to view the size of the `cdc` data frame; this will return the number of rows and columns. (In RStudio, you can see this information by simply looking next to the object name in the *Environment* pane in the upper right-hand corner).

### Subsetting (Base `R`)

```{r dim, exercise = TRUE}
`___`(`___`)
```

```{r dim-solution}
dim(cdc)
```

Now, if we want to access a subset of the full data frame, we can use row-and-column notation. For example, to see the sixth variable of the 567<sup>th</sup> respondent, use the format

```{r cdc-row567-column6, eval = FALSE, echo = TRUE}
cdc[567,6]
```

which means we want the element of our data set that is in the 567<sup>th</sup> row (meaning the 567<sup>th</sup> person or observation) and the 6<sup>th</sup> column (in this case, `weight`). We know that `weight` is the 6<sup>th</sup> variable because it is the 6<sup>th</sup> entry in the list of variable names.

```{r names-again, eval = TRUE, echo = TRUE}
names(cdc)
```

To see the weights for the first 10 respondents we can type

```{r first-10-rows-sixth-column, eval = FALSE, echo = TRUE}
cdc[1:10,6]
```

In this expression, we have asked just for rows in the range 1 through 10.  `R` uses the `:` to create a range of values, so 1:10 expands to 1, 2, 3, 4, 5, 6, 7, 8, 9, 10. 

Finally, if we want all of the data for the first 10 respondents, type

```{r first-10-rows, eval = FALSE, echo = TRUE}
cdc[1:10,]
```

By leaving out an index or a range (we didn't type anything between the comma and the square bracket), we get all the columns. When starting out in `R`, this is a bit counterintuitive. As a rule, we omit the column number to see all columns in a data frame. Similarly, if we leave out an index or range for the rows, we would access all the observations, not just the 567<sup>th</sup>, or rows 1 through 10. Display the weights of the second batch of 100 respondents.

```{r 6th-column, exercise = TRUE}
cdc[`___`, `___`]
```

```{r 6th-column-check}
grade_result(
  pass_if(~ .result[3] ==110),
  fail_if(~ length(.result) != 100, "This vector doesn't seem to be the correct size."),
  fail_if(~ .result[1] != 118, "This vector doesn't seem to start at the correct index.")
)
```

Recall that column 6 represents respondents' weight, so the command above reported all of the weights in the data set. An alternative method to access the weight data is by referring to the name.  Previously, we typed `names(cdc)` to see all the variables contained in the cdc data set. We can use any of the variable names to select items in our data set. 

```{r weight, eval = FALSE, echo = TRUE}
cdc$weight
```

The dollar-sign tells `R` to look in data frame `cdc` for the column called `weight`.  Since that's a single vector, we can subset it with just a single index inside square brackets.  We see the weight for the 567<sup>th</sup> respondent by typing

```{r weight-567, eval = FALSE, echo = TRUE}
cdc$weight[567]
```

Write a code to retrieve the weights of the first 10 respondents using the dollar-sign notation. 

```{r weight-first10, exercise = TRUE}
`___`
```

The command above returns the same result as the `cdc[1:10,6]` command. Both row-and-column notation and dollar-sign notation are widely used, which one you choose to use depends on your personal preference.

### A little more on subsetting (the `tidy` way)

It's often useful to extract all individuals (cases) in a data set that have specific characteristics. We accomplish this through *conditioning* commands. First, consider expressions like

```{r true-male, eval = FALSE, echo = TRUE}
cdc$gender == "m"
```

or

```{r true-over30, eval = FALSE, echo = TRUE}
cdc$age > 30
```

These commands produce a series of `TRUE` and `FALSE`  values. There is one value for each respondent, where `TRUE` indicates that the person was male (via the first command) or older than 30 (second command).

Suppose we want to extract just the data for the men in the sample, or just for those over 30. We can pipe the `cdc` data frame to the `filter` function to do that for us. For example, the command

<!-- ```{r males, exercise = TRUE} -->
<!-- mdata <- subset(cdc, gender == "m") -->
<!-- ``` -->

```{r males, exercise = TRUE}
mdata <- cdc %>% filter(gender == "m")
```

will create a new data set called `mdata` that contains only the men from the `cdc` data set. This new data set contains all the same variables but just under half the rows. The important thing is that we can carve up the data based on values of one or more variables.

:::{.noteblock}
**Logical operators:** Filtering for certain observations (e.g. response by gender or age) is often of interest in data frames where we might want to examine observations with certain characteristics separately from the rest of the data.
To do so, you can use the `filter` function and a series of **logical operators**.
The most commonly used logical operators for data analysis are as follows:

-   `==` means "equal to"
-   `!=` means "not equal to"
-   `>` or `<` means "greater than" or "less than"
-   `>=` or `<=` means "greater than or equal to" or "less than or equal to"
:::

Create a subset of the health data of men over thirty. 

```{r males-and-over30, exercise = TRUE}
m_and_over30 <- cdc %>% filter(gender == "m", age > 30)
m_and_over30 
```

```{r males-and-over30-check}
grade_result(
  pass_if(~ all((dim(.result) == c(7244, 9)) & (.result['age'] > 30) & (.result['gender'] == 'm'))),
  fail_if(~ !all(.result['age'] > 30), "The ages are not above 30."),
  fail_if(~ !all(.result['gender'] == "m"), "Not all of the respondents are male."),
  fail_if(~ !all(dim(.result) == c(7244, 9)), "This vector doesn't seem to be the correct size."),
)
```

Note that you can separate `filter` conditions using commas if you want responses that are both from males **and** over 30.
If you are interested in either men **or** over 30, then you can use the `|` instead of the comma.

Create a subset of the health data of mean or people over thirty. 

```{r males-or-over30, exercise = TRUE}
m_or_over30 <- cdc %>% filter(gender == "m" | age > 30)
m_or_over30
```

```{r males-or-over30-check}
grade_result(
  pass_if(~ all((dim(.result) == c(17591, 9)) & ((.result['age'] > 30) | (.result['gender'] == 'm')))),
  fail_if(~ !all(.result['age'] > 30), "The ages are not above 30."),
  fail_if(~ all(.result['gender'] == "m"), "All of the respondents are male. You should be females over thrity too."),
  fail_if(~ !all(dim(.result) == c(17591, 9)), "This vector doesn't seem to be the correct size.")
)
```

In principle, you may use as many "and" and "or" clauses as you like when filtering data.

:::{.exerciseblock}
Using the **CDC Health Data**, create a new object called `under23_and_smoke` that contains all observations of respondents under the age of 23 that have smoked 100 cigarettes in their lifetime. Write the command you used to create the new object as the answer to this exercise.

```{r answer-6, exercise = TRUE}
under23_and_smoke <- `___`
under23_and_smoke
```

```{r answer-6-solution}
under23_and_smoke <- cdc %>% filter(age < 23, smoke100 == 1)
under23_and_smoke
```
:::


## Quantitative data

With our subsetting tools in hand, we'll now return to the task of the day: making basic summaries of the **CDC Health Data**.  We've already looked at categorical data such as `smoke` and `gender` so now let's turn our attention to quantitative data.  Two common ways to visualize quantitative data are with box plots and histograms.  We can construct a box plot for a single variable with the following command.

```{r boxplot-height, exercise = TRUE}
ggplot(cdc, aes(x = height)) + geom_boxplot()
```

You can compare the locations of the components of the box by examining the summary statistics.

```{r summary-height, exercise = TRUE}
summary(cdc$height)
```

Confirm that the median and upper and lower quartiles reported in the numerical summary match those in the graph. The purpose of a box plot is to provide a thumbnail sketch of a variable for the purpose of comparing across several categories. We can do this with `ggplot` by adding a `y` aesthetic to the mix.

```{r boxplot-height-gender, exercise = TRUE}
ggplot(cdc, aes(x = height, y = gender)) + geom_boxplot()
```

Next let's consider a new variable that doesn't show up directly in this data set: Body Mass Index (BMI) [[W]](http://en.wikipedia.org/wiki/Body_mass_index). BMI is a weight to height ratio and can be calculated as:

\[ BMI = \frac{weight~(lb)}{height~(in)^2} * 703 \]

The factor of 703 is the approximate conversion to change units from metric (meters and kilograms) to imperial (inches and pounds).

To create a new variable called `bmi` we use the `mutate` function. Then we can create a box plot of `bmi` by `genhlth`.

```{r boxplot-bmi, exercise = TRUE}
cdc <- cdc %>% mutate(bmi = 703 * weight / height^2)
ggplot(cdc, aes(x = genhlth, y = bmi)) + geom_boxplot()
```

Notice that the first line above is just some arithmetic, but it's applied to all 20,000 numbers in the `cdc` data set. That is, for each of the 20,000 participants, we take their weight, divide by their height-squared and then multiply by 703. The result is 20,000 BMI values, one for each respondent. This is one reason why we like R: it lets us perform computations like this using very simple expressions.

:::{.exerciseblock}
Using the **CDC Health Data**, create a box plot of `bmi` versus `genhlth`. What does this box plot show? Pick another categorical variable from the data set and see how it relates to BMI. List the variable you chose, why you might think it would have a relationship to BMI,  and indicate what the figure seems to suggest. 

```{r answer-7a, exercise = TRUE}
cdc <- `___` %>% mutate(bmi = 703 * `___` / `___`)
ggplot(`___`) + `___`
```

```{r answer-7b}
question_text(
  "Enter your solution as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```    
:::

Finally, let's make some histograms using the command `hist`. Create a histogram for the age of our respondents.

```{r hist-age, exercise = TRUE}
ggplot(cdc, aes(x = age)) + geom_histogram()
```

Histograms are generally a very good way to see the shape of a single distribution, but that shape can change depending on how the data is split between the different bins.  You can control the number of bins by adding an argument to the command `geom_histogram`. Experiment with changing the number of `bin` (e.g., replace the 30 below with something smaller and then something bigger). 

```{r hist-bmi, exercise = TRUE}
ggplot(cdc, aes(x = age)) + geom_histogram(bins = 30)
```

How do these two histograms compare?

```{r hist-free}
question_text(
  "Enter your observations as free text below.",
  answer("C0rrect", correct = TRUE),
  incorrect = "Okay! [even if this button is red]",
  try_again_button = "Modify your answer",
  allow_retry = TRUE
)
```

At this point, we've done a good first pass at analyzing the information in the BRFSS questionnaire.  We've found an interesting association between smoking and gender, and we can say something about the relationship between people's assessment of their general health and their own BMI.  We've also picked up essential computing tools -- summary statistics, subsetting, and plots -- that will serve us well throughout this course.

## Summary (Exercises 1-6)

:::{.warningblock}
To submit your completed laboratory report, submit the answers (only) to the exercises 1-6. Upload the submission to [Gradescope](https://www.gradescope.com) as a PDF. Further instructions on using the class template to produce a nice lab report containing both data analysis and text are in the first Section "RStudio and lab reports" of Lab 1. 
:::

The following exercises refer to variables found in or derived from the **CDC Health Data**  in the data frame `cdc` which is packaged with `MA22004labs`. 

<!-- 1.  Using the **CDC Data**, create a numerical summary for `height` and `age`, and compute the interquartile range for each. Compute the relative frequency distribution for `gender` using `tabyl`. How many males are in the sample? What proportion of the sample reports being in excellent health?  **[5 marks]** -->

1.  Generate a mosaic plot of `gender` vs `smoke100` from the **CDC Health Data**. What does the mosaic plot of  reveal about smoking habits and gender? **[2 marks]**

<!-- 2.  Using the **CDC Health Data**, create a new object called `under23_and_smoke` that contains all observations of respondents under the age of 23 that have smoked 100 cigarettes in their lifetime. Write the command you used to create the new object as the answer to this exercise. **[1 mark]** -->

2.  Using the **CDC Health Data**, pick another categorical variable from `cdc` and see how it relates to BMI. List the variable you chose, why you might think it would have a relationship to BMI,  and indicate what the figure seems to suggest. **[4 marks]**

3.  Using the **CDC Health Data**, make a scatter plot of weight versus desired weight. Describe the relationship between these two variables. **[3 marks]**

4.  Using the **CDC Health Data**, consider a new variable: the difference between desired weight (`wtdesire`) and current weight (`weight`). Create this new variable by subtracting the two columns in the data frame and assigning them to a new object called `wdiff`. Write the command you used to create the new object as the answer to this exercise. What type of data is `wdiff`? If an observation `wdiff` is 0, what does this mean about the person's weight and desired weight? What if `wdiff` is positive or negative? **[3 marks]**

5.  By using appropriate numerical summaries and plots, describe the distribution of `wdiff`, created in Exercise 4, in terms of its center, shape, and spread, including any plots you use. What does this tell us about how people feel about their current weight? **[4 marks]**

6.  Using numerical summaries and a side-by-side box plot, determine if men tend to view their weight differently 
than women.  **[4 marks]**

<!-- 7.  Find the mean and standard deviation of `weight` in `cdc` and determine what proportion of the weights are within one standard deviation of the mean. Give a box plot for the values of `wdiff` for the subset of individuals whose weight is  within one standard deviation of the mean weight. How does the distribution of `wdiff` vary for this subset? **[6 marks]** -->


------------------------------------------------------------------------

<a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">![Creative Commons License](https://i.creativecommons.org/l/by-sa/4.0/88x31.png){style="border-width:0"}</a><br />This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a>.

This work is derivative of OpenIntro Labs and was modified by Eric Hall (2020). The original labs were adapted for OpenIntro by Andrew Bray and Mine Çetinkaya-Rundel from labs written by Mark Hansen of UCLA Statistics.

